```python
# Initialize Otter
import otter
grader = otter.Notebook("hw07.ipynb")
```

# Homework 7: Confidence Intervals

<img src="data8logo.png" alt="Data 8 Logo" style="width: 15%; float: right; padding: 1%; margin-right: 2%;"/>

**Helpful Resource:**

- [Python Reference](http://data8.org/sp26/reference/): Cheat sheet of helpful array & table methods used in this data science course!

**Recommended Reading**: 

* [Estimation](https://www.inferentialthinking.com/chapters/13/Estimation)

Please complete this notebook by filling in the cells provided. **Before you begin, execute the cell below to setup the notebook by importing some helpful libraries.** Each time you start your server, you will need to execute this cell again.

For all problems that you must write explanations and sentences for, you **must** provide your answer in the designated space. Moreover, throughout this homework and all future ones, **please be sure to not re-assign variables throughout the notebook!** For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!

Directly sharing answers is not okay, but discussing problems with the course staff or with other students is encouraged.

You should start early so that you have time to get help if you're stuck.

---

The point breakdown for this assignment is given in the table below:
| Category | Points |
| --- | --- |
| Autograder (Coding questions) | 92 |
| Written (Visualization questions) | 8 |
| **Total** | 100 |

```python
# Don't change this cell; just run it. 

import numpy as np
from datascience import *

# These lines do some fancy plotting magic.",
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt 
plt.style.use('fivethirtyeight')
import warnings
warnings.simplefilter('ignore', FutureWarning)
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

## 1. Italian Restaurants in Arcata/Eureka

Bonnie and Bethany are trying to see what the best Italian restaurant is in the area. They survey 1,500 Cal Poly Humboldt students selected uniformly at random and ask each student what Italian restaurant is the best. (*Note: This data is fabricated for the purposes of this homework.*) The choices of Italian restaurants are [Venicia Italian Restaurant](https://www.google.com/maps/place/Mazzotti's+-+Italian+Restaurant/@40.8681176,-124.0868234,15z/data=!4m6!3m5!1s0x54d1570f63be6e0b:0x55e3986e9f3226f4!8m2!3d40.8681176!4d-124.0868234!16s%2Fg%2F1v27v6zr?entry=ttu), [Brick & Fire](https://www.google.com/maps/place/Brick+%26+Fire/@40.7923298,-124.166398,17z/data=!3m1!4b1!4m6!3m5!1s0x54d3ffddc91d64df:0x8a806dade51653e9!8m2!3d40.7923258!4d-124.1638231!16s%2Fg%2F11f55d5rg7?entry=ttu), [Babe's Pizza & Pasta](https://www.google.com/maps/place/Babe's+Pizza+%26+Pasta/@40.7694284,-124.145214,17z/data=!3m1!4b1!4m6!3m5!1s0x54d3ff9e61a3a4ad:0xb708b698daf884b!8m2!3d40.7694244!4d-124.1426391!16s%2Fg%2F1tf6lyln?entry=ttu), and [Gabriel's Italian Restaurant](https://www.google.com/maps/place/Gabriel%E2%80%99s+Italian+Restaurant/@40.8038929,-124.1703527,17z/data=!3m1!4b1!4m6!3m5!1s0x54d1556528cd9149:0xf35a84f0bb853f90!8m2!3d40.8038889!4d-124.1677778!16s%2Fg%2F1tqg01jw?entry=ttu). After compiling the results, Bonnie and Bethany release the following percentages from their sample:

|Italian Restaurant  | Percentage|
|:------------:|:------------:|
|Venicia | 8% |
|Brick & Fire | 53% |
|Babe's Pasta & Pizza | 25% |
|Gabriel's | 14% |

These percentages represent a uniform random sample of the population of Cal Poly Humboldt students. We will attempt to estimate the corresponding *parameters*, or the percentage of the votes that each restaurant would receive from the population (i.e. all Cal Poly Humboldt students). We will use confidence intervals to compute a range of values that reflects the uncertainty of our estimates.

The table `votes` contains the results of Bonnie and Bethany's survey.

```python
# Just run this cell
votes = Table.read_table('votes.csv')
votes
```

---

**Question 1.1.** Complete the function `one_resampled_percentage` below. It should return Brick & Fire's ***percentage*** of votes after taking the original table (`tbl`) and performing one bootstrap sample of it. Remember that a percentage is between 0 and 100. **(8 Points)**

*Note 1:* `tbl` will always be in the same format as `votes`.

*Note 2:* This function should be completed **without** `.group` or `.pivot`. Using these functions will cause your code to timeout.

*Hint:* Given a table of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**


```python
def one_resampled_percentage(tbl):
    ...

one_resampled_percentage(votes)
```

```python
grader.check("q1_1")
```

---

**Question 1.2.** Complete the `percentages_in_resamples` function such that it simulates and returns an array of **2025 elements**, where each element represents a bootstrapped estimate of the percentage of voters who will vote for Brick & Fire. You should use the `one_resampled_percentage` function you wrote above. **(8 Points)**

_Note:_ We perform our simulation with **only 2025 trials** in this problem to reduce the runtime, but we should generally use more repetitions.


```python
def percentages_in_resamples():
    percentage_brick_fire = make_array()
    ...
```

```python
grader.check("q1_2")
```

In the following cell, we run the function you just defined, `percentages_in_resamples`, and create a histogram of the calculated statistic for the 2025 bootstrap estimates of the percentage of voters who voted for Brick & Fire. 

*Note:* This might take a few seconds to run.

```python
resampled_percentages = percentages_in_resamples()
Table().with_column('Estimated Percentage', resampled_percentages).hist("Estimated Percentage")
```

---

**Question 1.3.** Using the array `resampled_percentages`, find the values at the two ends of the middle 95% of the bootstrapped percentage estimates. Assign the lower and upper ends of the interval to `brick_fire_lower_bound` and `brick_fire_upper_bound` respectively. **(8 Points)**

*Hint:* If you are stuck on this question, try looking over [Chapter 13.1](https://inferentialthinking.com/chapters/13/1/Percentiles.html) of the textbook, and check out the `percentile` function on the [python reference sheet](https://www.data8.org/sp26/reference/).


```python
brick_fire_lower_bound = ...
brick_fire_upper_bound = ...
print(f"Bootstrapped 95% confidence interval for the percentage of Brick & Fire voters in the population: [{brick_fire_lower_bound:.2f}, {brick_fire_upper_bound:.2f}]")
```

```python
grader.check("q1_3")
```

---

**Question 1.4.** The survey results seem to indicate that Brick & Fire is beating all the other Italian restaurants among the voters. We would like to use confidence intervals to determine a range of likely values for Brick & Fire's percentage lead over all the other restaurants combined. The calculation for Brick & Fire's lead over Venicia, Babe's, and Gabriel's combined is:

$$\text{Brick & Fire's \% of Vote} - (\text{100\%} - \text{Brick & Fire's \% of Vote})$$

Define the function `one_resampled_difference` that returns **exactly one value** of Brick & Fire's percentage lead over Venicia, Babe's, and Gabriel's combined from one bootstrap sample of `tbl`. **(8 Points)**

*Hint 1:* Brick & Fire's lead can be negative.

*Hint 2:* Given a table of votes, how can you figure out what percentage of the votes are for a certain restaurant? **Be sure to use percentages, not proportions, for this question!**

*Note:* If the skeleton code provided within the function is not helpful for you, feel free to approach the question using your own variables.


```python
def one_resampled_difference(tbl):
    bootstrap = ...
    brick_fire_percentage = ...
    ...
```

```python
grader.check("q1_4")
```

<!-- BEGIN QUESTION -->

---

**Question 1.5.** Write a function called `leads_in_resamples` that computes 2025 bootstrapped elements of Brick & Fire's percentage lead over Venicia, Babe's, and Gabriel's combined (the result of calling `one_resampled_difference`). It should return an array of the 2025 bootstrapped estimates. Afterwards, run the cell to plot a histogram of the resulting samples. **(8 Points)**

*Hint:* If you see an error involving `NoneType`, consider what components a function needs to have!


```python
def leads_in_resamples():
    ...

sampled_leads = leads_in_resamples()
Table().with_column('Estimated Lead', sampled_leads).hist("Estimated Lead")
```

<!-- END QUESTION -->

---

**Question 1.6.** Use the simulated data in `sampled_leads` from Question 1.5 to compute an approximate 95% confidence interval for Brick & Fire's percentage lead over Venicia, Babe's, and Gabriel's combined. **(10 Points)**


```python
diff_lower_bound = ...
diff_upper_bound = ...
print("Bootstrapped 95% confidence interval for Brick & Fire's true lead over Venicia, Babe's, and Gabriel's combined: [{:f}%, {:f}%]".format(diff_lower_bound, diff_upper_bound))
```

```python
grader.check("q1_6")
```

<hr style="border: 5px solid #003262;" />
<hr style="border: 1px solid #fdb515;" />

### 2. Interpreting Confidence Intervals 

Dr. Johnson computed the following 95% confidence interval for the percentage of Brick & Fire voters: 

$$[50.53, 55.53]$$

(Your answer from 1.3 may have been a bit different due to randomness; that doesn't mean it was wrong!)

---

**Question 2.1.** Dr. Johnson also created 70%, 90%, and 99% confidence intervals from the same sample, but he forgot to label which confidence interval represented which percentages! ***First***, **match each confidence level (70%, 90%, 99%) with one of the corresponding intervals given** (e.g. `CI_70_percent = 1`). **(5 points)**

The intervals are below:

1. [50.03, 55.94]
2. [52.1, 54]
3. [50.97, 54.99]

*Hint:* If you are stuck on this question, try looking over Chapters [13.3](https://inferentialthinking.com/chapters/13/3/Confidence_Intervals.html) and [13.4](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook.


```python
CI_70_percent = ...
CI_90_percent = ...
CI_99_percent = ...
```

```python
grader.check("q2_1")
```

**Question 2.2.** How did you arrive at your answer to Question 2.1? Select all correct answers that correctly interpet confidence levels and assign them to the array `confidence_answers`. (e.g. `confidence_answers = make_array(1,2,3)`) **(5 Points)**

1. Higher confidence levels lead to **wider** intervals because we include **more** of the statistics.

2. As confidence increases, we become **more** confident that the population parameter lies within our confidence interval.

3. Higher confidence levels lead to **smaller** intervals because we include **less** of the statistics.

4. As confidence increases, we become **less** confident that the population parameter lies within our confidence interval.

```python
confidence_answers = make_array(...)
```

```python
grader.check("q2_2")
```

---

**Question 2.3.** Suppose we produced **6,000 new samples** (each one a new/distinct uniform random sample of 1,500 students) from the population and **created a 95% confidence interval from each one**. **Roughly how many of those 6,000 intervals do you expect will actually contain the true percentage of the population?** **(10 Points)**

Assign your answer to `true_percentage_intervals`.


```python
true_percentage_intervals = ...
```

```python
grader.check("q2_3")
```

Recall the second bootstrap confidence interval you created, which estimated Brick & Fire's lead over Venicia, Babe's, and Gabriel's combined. Among
voters in the sample, Brick & Fire's lead was 6%. Dr. Johnson's 95% confidence interval for the true lead (in the population of all voters) was:

$$[1.2, 11.2]$$

Suppose we are interested in testing a simple yes-or-no question:

> "Is the percentage of votes for Brick & Fire equal to the percentage of votes for Venicia, Babe's, and Gabriel's combined?"

Our null hypothesis is that the percentages are equal, or equivalently, that Brick & Fire's lead is exactly 0. Our alternative hypothesis is that Brick & Fire's lead is not equal to 0.  In the questions below, don't compute any confidence interval yourself—use only Dr. Johnson's 95% confidence interval.

> Hint: Try thinking about the width of the 95% confidence interval in comparison to the new confidence intervals in the questions below. **Drawing a picture may help.**

---

**Question 2.4.** Say we use a 5% p-value cutoff. Do we reject the null, fail to reject the null, or are we unable to tell using Dr. Johnson's confidence interval? **(10 Points)**

Assign `cutoff_five_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using Dr. Johnson's confidence interval

*Hint:* Consider the relationship between the p-value cutoff and confidence. If you're confused, take a look at [this chapter](https://inferentialthinking.com/chapters/13/4/Using_Confidence_Intervals.html) of the textbook.


```python
cutoff_five_percent = ...
```

```python
grader.check("q2_4")
```

---

**Question 2.5.** What if, instead, we use a p-value cutoff of 1%? Do we reject the null, fail to reject the null, or are we unable to tell using Dr. Johnson's confidence interval? **(10 Points)**

Assign `cutoff_one_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using Dr. Johnson's confidence interval


```python
cutoff_one_percent = ...
```

```python
grader.check("q2_5")
```

---

**Question 2.6.** What if we use a p-value cutoff of 10%? Do we reject, fail to reject, or are we unable to tell using our confidence interval? **(10 Points)**

Assign `cutoff_ten_percent` to the number corresponding to the correct answer.

1. Reject the null / Data is consistent with the alternative hypothesis
2. Fail to reject the null / Data is consistent with the null hypothesis
3. Unable to tell using Dr. Johnson's confidence interval


```python
cutoff_ten_percent = ...
```

```python
grader.check("q2_6")
```

You're done with Homework 7!  

**Important submission steps:** 
1. Run the tests and verify that they all pass.
2. Choose **Save Notebook** from the **File** menu, then **run the final cell**. 
3. Click the link to download the zip file.
4. Go to your course's submission platform and submit the zip file to the corresponding assignment. The name of this assignment is "HW 07 Autograder". 

**It is your responsibility to make sure your work is saved before running the last cell.**

### Pets of this data science course

**Harbor** and **Loki** (Joel's) are proud of you for completing the assignment!

<img src="f26_harbor_and_loki.jpg" width="50%" alt="Joel's pet Harbor and Loki"/>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

## Submission

Make sure you have run all cells in your notebook in order before running the cell below, so that all images/graphs appear in the output. The cell below will generate a zip file for you to submit. **Please save before exporting!**

```python
# Save your notebook first, then run this cell to export your submission.
grader.export(pdf=False)
```

Tests for q1_1 (points: [0, 0, 2, 6])
```
>>> bool(type(one_resampled_percentage(votes)) in set([float, np.float64]))
True
```
```
>>> bool(one_resampled_percentage(votes) > 1)
True
```
```
>>> bool(35 <= one_resampled_percentage(votes) <= 65)
True
```
```
>>> from hw07_check import correctness_check_1_1
>>> correctness_check_1_1(one_resampled_percentage, votes)
True
```


Tests for q1_2 (points: [2, 2, 4])
```
>>> bool(len(percentages_in_resamples()) == 2025)
True
```
```
>>> from hw07_check import correctness_check_1_2_1
>>> correctness_check_1_2_1(percentages_in_resamples)
True
```
```
>>> from hw07_check import correctness_check_1_2_2
>>> correctness_check_1_2_2(percentages_in_resamples)
True
```


Tests for q1_3 (points: [0, 8])
```
>>> bool(45 <= brick_fire_lower_bound <= brick_fire_upper_bound <= 65)
True
```
```
>>> from hw07_check import correctness_check_1_3
>>> correctness_check_1_3(brick_fire_lower_bound, brick_fire_upper_bound, resampled_percentages)
True
```


Tests for q1_4 (points: [0, 0, 8])
```
>>> bool(type(one_resampled_difference(votes)) in set([float, np.float64]))
True
```
```
>>> bool(0 <= one_resampled_difference(votes) <= 100)
True
```
```
>>> from hw07_check import correctness_check_1_4
>>> correctness_check_1_4(one_resampled_difference, votes)
True
```


Tests for q1_6 (points: [0, 2, 8])
```
>>> bool(-5 <= diff_lower_bound <= diff_upper_bound <= 20)
True
```
```
>>> from hw07_check import correctness_check_1_6_1
>>> correctness_check_1_6_1(diff_lower_bound, diff_upper_bound)
True
```
```
>>> from hw07_check import correctness_check_1_6_2
>>> correctness_check_1_6_2(diff_lower_bound, diff_upper_bound, sampled_leads)
True
```


Tests for q2_1 (points: [5])
```
>>> from hw07_check import correctness_check_2_1
>>> correctness_check_2_1(CI_70_percent, CI_90_percent, CI_99_percent)
True
```


Tests for q2_2 (points: [5])
```
>>> from hw07_check import correctness_check_2_2
>>> correctness_check_2_2(confidence_answers)
True
```


Tests for q2_3 (points: [0, 10])
```
>>> bool(1000 <= true_percentage_intervals <= 10000)
True
```
```
>>> from hw07_check import correctness_check_2_3
>>> correctness_check_2_3(true_percentage_intervals)
True
```


Tests for q2_4 (points: [0, 0, 10])
```
>>> bool(1 <= cutoff_five_percent <= 3)
True
```
```
>>> bool(type(cutoff_five_percent) == int)
True
```
```
>>> from hw07_check import correctness_check_2_4
>>> correctness_check_2_4(cutoff_five_percent)
True
```


Tests for q2_5 (points: [0, 10])
```
>>> bool(1 <= cutoff_one_percent <= 3)
True
```
```
>>> from hw07_check import correctness_check_2_5
>>> correctness_check_2_5(cutoff_one_percent)
True
```


Tests for q2_6 (points: [0, 10])
```
>>> bool(1 <= cutoff_ten_percent <= 3)
True
```
```
>>> from hw07_check import correctness_check_2_6
>>> correctness_check_2_6(cutoff_ten_percent)
True
```
